Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Abubakar Usman"
COLLABORATORS = ""
```

---

# Assignment 1 - Python
(15 points total)

## Introduction

## Learning Outcomes
In this assignment you will practice:

- Basic Python syntax
- Writing simple Python programs

## Preparation

### Question 1
(3 points)

<!-- grade_id: cell-fb83ae036d49cd54, points: 0, grade: false, solution: true -->
```python
# Write a function that returns True if a point is within a rectangle specified by a top left
# corner and a bottom right corner
# e.g. isIn((1,2), (3,4), (1.5, 3.2)) should return True
# isIn((-1,0), (5,5), (6,0)) should return False
def isIn(topLeft=(0,0), bottomRight=(0,0), point=(0,0)):
    # YOUR CODE HERE
    if point[0] >= topLeft[0] and point[0] <= bottomRight[0] and point[1] >= topLeft[1] and point[1] <= bottomRight[1]:
        return True
    else:
        return False
    raise NotImplementedError()
```

**Question 1 Rubric.**
(3 points)
The answer to the main question 3 points,
1 point per upheld property (i.e., when `isIn(...)` holds, doesn't hold, and a basic case with the coordinate origin).

```python
isIn((1,2), (3,4), (1.5, 3.2))
```

Output:
```
True
```

```python
isIn((-1,0), (5,5), (6,0)) 
```

Output:
```
False
```

```python
isIn((0,0), (0,0), (0,0)) 
```

Output:
```
True
```

<!-- grade_id: cell-5ce20dc0f9c72478, points: 1, grade: true, solution: false -->
```python
assert isIn((1,2), (3,4), (1.5, 3.2))
```

<!-- grade_id: cell-6230439af6632d62, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-281973bd9194a0c1, points: 1, grade: true, solution: false -->
```python

```

### Question 2
(3 points)

<!-- grade_id: cell-4de3ad07063b602f, points: 0, grade: false, solution: true -->
```python
# Modify your function so it takes a list of points rather than a single point and returns True only if all points in the list
# are in the rectangle
# e.g. allIn((0,0), (5,5), [(1,1), (0,0), (5,5)]) should return True
# allIn((0,0), (5,5), [(1,1), (0,0), (5,6)]) should return False
def allIn(topLeft=(0,0), bottomRight=(0,0), pointList=[]):
    # YOUR CODE HERE
    if not pointList:
        return 
    for point in pointList:
        result = isIn(topLeft, bottomRight, point)
        if result == False:
            return False
    return True
    raise NotImplementedError()
```

**Question 2 Rubric.**
(3 points)
The answer to the main question 3 points,
1 point per upheld property (i.e., when `allIn(...)` holds, doesn't hold, and a basic case with the coordinate origin).

```python
allIn((0,0), (5,5), [(1,1), (0,0), (5,5)])
```

Output:
```
True
```

```python
allIn((0,0), (5,5), [(1,1), (0,0), (5,6)])
```

Output:
```
False
```

```python
allIn((0,0), (0,0), [(0,0), (0,0), (0,0)])
```

Output:
```
True
```

```python
allIn((0,0), (0,0), [])
```

<!-- grade_id: cell-93d6a064af509b2d, points: 1, grade: true, solution: false -->
```python
assert not allIn((0,0), (0,0), [])
```

<!-- grade_id: cell-facb252022b573af, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-fab05eae5cd85b29, points: 1, grade: true, solution: false -->
```python

```

### Question 3
(3 point)

In this question, you are given a string `s` which represents a DNA string. The string `s` consists of symbols 'A', 'C', 'G', and 'T'. An example of a length 21 DNA string is `"ATGCTTCAGAAAGGTCTTACG."` 

Your task is to write a code which will count the number of times each of the symbols 'A', 'C', 'G', and 'T' occur in `s`. The output should be a list of 4 integers.

```python
# Here is the DNA string:

s = 'AGCTTTTCATTCTGACTGCAACGGGCAATATGTCTCTGTGTGGATTAAAAAAAGAGTGTCTGATAGCAGC'
```

<!-- grade_id: cell-6e63758a67a1aef1, points: 0, grade: false, solution: true -->
```python
# Write the code to print out a list of four integers, which will represent the count
# of each of the symbols: 'A', 'C', 'G', and 'T'
def question3Function(input = s):
    # YOUR CODE HERE
    a, c, g, t = 0,0,0,0
    for item in input:
        if item == 'A':
            a += 1
        elif item == 'C':
            c += 1
        elif item == 'G':
            g += 1
        else:
            t += 1
    return [a, c, g, t]
    raise NotImplementedError()
```

**Question 3 Rubric.** (3 point) The exact answer to the main question is 3 points.

```python
question3Function()
```

Output:
```
[20, 12, 17, 21]
```

<!-- grade_id: cell-f6eb1bf9545aa152, points: 3, grade: true, solution: false -->
```python

```

### Question 4
(6 point)

```python
# Run this cell to create a dictionary of states' capitals
capitals={
    'Alabama': 'Montgomery',
    'Alaska': 'Juneau',
    'Arizona':'Phoenix',
    'Arkansas':'Little Rock',
    'California': 'Sacramento',
    'Colorado':'Denver',
    'Connecticut':'Hartford',
    'Delaware':'Dover',
    'Florida': 'Tallahassee',
    'Georgia': 'Atlanta',
    'Hawaii': 'Honolulu',
    'Idaho': 'Boise',
    'Illinios': 'Springfield',
    'Indiana': 'Indianapolis',
    'Iowa': 'Des Monies',
    'Kansas': 'Topeka',
    'Kentucky': 'Frankfort',
    'Louisiana': 'Baton Rouge',
    'Maine': 'Augusta',
    'Maryland': 'Annapolis',
    'Massachusetts': 'Boston',
    'Michigan': 'Lansing',
    'Minnesota': 'St. Paul',
    'Mississippi': 'Jackson',
    'Missouri': 'Jefferson City',
    'Montana': 'Helena',
    'Nebraska': 'Lincoln',
    'Neveda': 'Carson City',
    'New Hampshire': 'Concord',
    'New Jersey': 'Trenton',
    'New Mexico': 'Santa Fe',
    'New York': 'Albany',
    'North Carolina': 'Raleigh',
    'North Dakota': 'Bismarck',
    'Ohio': 'Columbus',
    'Oklahoma': 'Oklahoma City',
    'Oregon': 'Salem',
    'Pennsylvania': 'Harrisburg',
    'Rhoda Island': 'Providence',
    'South Carolina': 'Columbia',
    'South Dakota': 'Pierre',
    'Tennessee': 'Nashville',
    'Texas': 'Austin',
    'Utah': 'Salt Lake City',
    'Vermont': 'Montpelier',
    'Virginia': 'Richmond',
    'Washington': 'Olympia',
    'West Virginia': 'Charleston',
    'Wisconsin': 'Madison',
    'Wyoming': 'Cheyenne'  
}
# Write a function to return a list of all capitals that contain the name of a state in their name as a substring
# The list that is produced should include Kansas City and one other
```

**Question 4 Rubric.** (6 point) The order independent answer to the main question is 6 points.

<!-- grade_id: cell-c77e2c4238016c36, points: 0, grade: false, solution: true -->
```python
def state_in_capital_name(capitals):
    # YOUR CODE HERE
    my_list = []
    for key, value in capitals.items():
        if key in value:
            my_list.append(value)
    return my_list
    raise NotImplementedError()
```

```python
state_in_capital_name(capitals)
```

Output:
```
['Indianapolis', 'Oklahoma City']
```

<!-- grade_id: cell-8fb3410b996d49c3, points: 6, grade: true, solution: false -->
```python

```